# Bibliothèque de graphes

Le fichier [graph.py](./graph.py) contient le squelette de la classe Graph que vous avez implanté la séance dernière. Nous y avons depuis ajouté des tests. Éditez ce squelette pour y insérer le code de chacune des méthodes. Veillez à bien conserver tous les tests fournis! N'hésitez pas à compléter ces tests.
Essayez votre classe:

```python
from graph import Graph, examples
```

<!-- grade_id: cell-1dd035627b5d7b54, points: 1, grade: true, solution: false -->
```python
G = Graph([1,2,3], edges=[(1,2,"12"),(2,3,"23")])
```

<!-- grade_id: cell-f32bdde5da338aed, points: 1, grade: true, solution: false -->
```python
assert G.edges() == ((1, 2, '12'), (2, 1, '12'), (2, 3, '23'), (3, 2, '23'))
```

<!-- grade_id: cell-50e1b14caae52140, points: 1, grade: true, solution: false -->
```python
G.show()
```

Output:
```
Figure(fig_margin={'top': 60, 'bottom': 60, 'left': 60, 'right': 60}, layout=Layout(height='400px', width='400…
```

<!-- grade_id: cell-d181f9b9b27b9a76, points: 1, grade: true, solution: false -->
```python
G = examples.cours_1_G()
G.edges()
```

Output:
```
((0, 2, 1),
 (1, 2, 1),
 (1, 5, 1),
 (2, 0, 1),
 (2, 1, 1),
 (2, 3, 1),
 (3, 2, 1),
 (3, 5, 1),
 (3, 4, 1),
 (4, 3, 1),
 (5, 1, 1),
 (5, 3, 1))
```

<!-- grade_id: cell-120e2364a899a12b, points: 1, grade: true, solution: false -->
```python
assert G.vertex_number() == 6
```

```python
G.edges()
```

Output:
```
((0, 2, 1),
 (1, 2, 1),
 (1, 5, 1),
 (2, 0, 1),
 (2, 1, 1),
 (2, 3, 1),
 (3, 2, 1),
 (3, 5, 1),
 (3, 4, 1),
 (4, 3, 1),
 (5, 1, 1),
 (5, 3, 1))
```

Maintenant, lancez tous les tests inclus dans la documentation des fonctions:

```python
!pytest --doctest-modules graph.py
```

Output:
```
[1m============================= test session starts ==============================[0m
platform linux -- Python 3.7.6, pytest-5.4.3, py-1.8.1, pluggy-0.13.1
rootdir: /home/jovyan/Algorithme_Avancee/TP2/ParcoursDeGraphes
collected 6 items                                                              [0m

graph.py [32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m                                                          [100%][0m

[32m============================== [32m[1m6 passed[0m[32m in 0.04s[0m[32m ===============================[0m
```

```python
G = examples.directed()
sorted(G.edges())
```

Output:
```
[(1, 2, 12), (1, 4, 12), (2, 3, 23)]
```

```python
G.matrix()
```

Output:
```
[[0, 12, 0, 12], [0, 0, 23, 0], [0, 0, 0, 0], [0, 0, 0, 0]]
```

```python

```