## [Parcours en largeur](https://fr.wikipedia.org/wiki/Algorithme_de_parcours_en_largeur) et calcul de distance

Dans cette feuille, nous allons rafiner l'agorithme de parcours de graphe vu précédemment, pour calculer des *distances simples* entre sommets d'un graphe, ne tenant pas compte des poids des arêtes. Nous suivrons la même démarche que pour notre premier algorithme: invariants, test sur des exemples, visualisation, complexité, preuve de correction.

**Définition** *distance simple* dans un graphe

La *distance* entre deux sommets $u$ et $v$ d'un graphe $G$ est le plus petit entier $l$ tel qu'il existe un chemin avec $l$ arêtes allant de $u$ à $v$. S'il n'y a pas de tel chemin, alors la distance est $\infty$.

**Exercice**

1. Complétez la fonction suivante qui implante un parcours en largeur, en vous laissant guider par les invariants fournis

<!-- grade_id: cell-fa2ebcc093dd410e, points: 0, grade: false, solution: true -->
```python
def parcours_en_largeur(G, u):
    """
    INPUT:
    - 'G' - un graphe
    - 'u' - un sommet du graphe
    
    OUTPUT: un dictionnaire associant à chaque sommet `v` accessible depuis `u` sa distance depuis `u`
    """
    distances = {u: 0} # L'ensemble des sommets déjà rencontrés
    todo      = [u]    # Une liste de sommets à traiter

    while todo:
                
        v = todo.pop()
        for w in G.neighbors_out(v):
            if w not in distances:
                distances[w] = distances[v] + 1
                todo.append(w)

        # Invariants:
        # - Si `v` est dans `distance`, alors il y a un chemin de `u` à `v`,
        #   et distance[v] contient la distance de `u` à `v`;
        # - Si `v` est dans `distance` et pas dans `todo`
        #   alors tous les voisins de `v` sont dans dans `distance`
        # YOUR CODE HERE
    return distances
```

2. Testez que votre fonction est correcte sur les exemples suivants:

```python
from graph import Graph, examples
```

```python
C3 = examples.C3()
C3.edges()
```

Output:
```
((0, 1, 1), (1, 2, 1), (2, 0, 1))
```

```python
parcours_en_largeur(C3, 0)
```

Output:
```
{0: 0, 1: 1, 2: 2}
```

<!-- grade_id: cell-e31ac5a2ec1ca694, points: 1, grade: true, solution: false -->
```python
assert parcours_en_largeur(C3, 0) == {0: 0, 1: 1, 2: 2}
```

```python
T3 = examples.T3()
T3.edges()
```

Output:
```
((0, 1, 1), (0, 2, 1), (1, 2, 1))
```

<!-- grade_id: cell-2a1fe26cce698ead, points: 1, grade: true, solution: false -->
```python
assert parcours_en_largeur(T3, 0) == {0: 0, 1: 1, 2: 1}
```

```python
from graph import examples
G = examples.parcours_directed()
```

<!-- grade_id: cell-8d83da87d56c871c, points: 1, grade: true, solution: false -->
```python
assert parcours_en_largeur(G, "H") == {'H': 0, 'F': 1, 'G': 2}
```

<!-- grade_id: cell-1cefa7a59eae7cc6, points: 1, grade: true, solution: false -->
```python
assert parcours_en_largeur(G, "A") == {'A': 0, 'B': 1, 'G': 1, 'F': 1, 'C': 2, 'H': 2, 'D': 3}
```

3. Instrumentez votre code pour visualiser son exécution

```python
def parcours_visualisation(G, u):
    """
    INPUT:
    - 'G' - un graphe
    - 'u' - un sommet du graphe
    
    OUTPUT: la liste des sommets `v` de `G`
            tels qu'il existe un chemin de `u` à `v`
    """
    distances = {u: 0} # L'ensemble des sommets déjà rencontrés
    todo = [u]    # Une liste de sommets à traiter
    player.player.reset(copy.deepcopy(locals()))
    while todo:
        # Invariants:
        # - Si `v` est dans `distance`, alors il y a un chemin de `u` à `v`,
        #   et distance[v] contient la distance de `u` à `v`;
        # - Si `v` est dans `distance` et pas dans `todo`
        #   alors tous les voisins de `v` sont dans dans `distance`
        v = todo.pop()
        # Observation des variables locales

        player.set_value(copy.deepcopy(locals()))
        for w in G.neighbors_out(v):
            if w not in distances:
                player.set_value(copy.deepcopy(locals()))
                distances[w] = distances[v] + 1
                todo.append(w)
        v = None
        player.set_value(copy.deepcopy(locals()))
        
    return distances 
```

**Exercice** Complexité

Donner une borne de complexité pour l'algorithme `parcours_en_largeur`.

Note: pour simplifier, nous avons utilisé une liste pour `todo`. Pour avoir une bonne complexité, il faudrait utiliser une *file* (fifo), typiquement implantée au moyen d'une liste chaînée et non un tableau comme dans les listes Python. Voir par exemple `queue.SimpleQueue`.

<!-- grade_id: cell-1a5e4461ec87ced1, points: 2, grade: true, solution: true -->
`un`=nombre de sommet u
<br>
`vn`=nombre d'arc 
<br>
`G`=graphe donné en parametre
<br>
`SO`=sommet original est le sommet donné en argument duquel nous cherchons la liste des plus courts chemins  des sommets `v` de `G`
            de`SO` à `v`
<br>
`todo`=Les sommets courants, sont les sommets auquel nous verifions tous les voisins
<br>
`C1`=Condition :Tant que tous les sommets ne sont pas marqués
<br>
Calculons $O(n)=$ complexité du pire cas
<br>
Tant que tous les sommets ne sont pas marqués :  dans le pire cas $O(un)$ (car un sommet est marqué au plus une fois):
<br>
    <li>nous supprimons le sommet courant de la liste : $O(1)$</li>
    <li>nous parcourons tous les sommets voisins du sommet courant excepter les voisins déjà visité (+1 car dans le premier cas, dans le pire des cas nous visitons tous les voisins sans qu'aucun autre sommet ne soit marqué):  dans le pire cas $O(vn/un+1)$</li>
    <li>Nous ajoutons la distance de tous les sommets voisins non marqué $O(1)$</li>
    <li>Nous ajoutons tous les sommets voisins non marqué a `todo`$O(1)$</li>
<br>
donc nous avons pour complexité $n=un*(1*((vn/un)+1)*(1+1))$ d'ou $O(n)=vn+un$

Bonus: chronométrez votre code sur des graphes de taille croissante afin de tracer une courbe de complexité pratique et vérifier empiriquement que votre code a la complexité voulue.

**Indication**: vous pouvez par exemple utiliser les outils `time.time`, `%timit`, `timeit` ou [bleachermark](https://github.com/miguelmarco/bleachermark).

**Exercice** Preuve de correction


<!-- grade_id: cell-f5da855a993a279c, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Conclusion

Dans cette feuille, vous avez mené en toute autonomie l'implantation et l'étude d'un autre parcours de graphe pour mettre en pratique tout ce que nous avions vu dans la fiche précédente. Bravo!

```python

```