## Plus courts chemins, avec poids: l'[algorithme de Dijkstra](https://fr.wikipedia.org/wiki/Algorithme_de_Dijkstra)

On considère le graphe suivant qui modélise un réseau routier:
<!-- By HB (Own work) [GFDL (http://www.gnu.org/copyleft/fdl.html) or CC BY-SA 3.0 (https://creativecommons.org/licenses/by-sa/3.0)], via Wikimedia Commons !-->
<img src="DijkstraBis01.svg">

On souhaite calculer le plus court chemin entre deux sommets de ce graphe.

**Exercice**
1. Implantez l'algorithme de Dijskstra

```python
from graph import Graph, examples
C3 = examples.C3()
d = {s : float("inf")  for s in C3.vertices()}
d
current_vertex = min(
                C3.vertices(), key=lambda vertex: d[vertex])

current_vertex
```

Output:
```
0
```

```python
C3.vertices(),d
key=lambda vertex: d[vertex]
min(C3.vertices(),key=lambda vertex: d[vertex])
```

Output:
```
0
```

<!-- grade_id: cell-b31db819a71cdffd, points: 3, grade: true, solution: true -->
```python
def dijkstra(G, e):
    """
    Renvoie un dictionnaire associant à chaque sommet sa distance depuis `e`
    """
    sousgraphe=list(G.vertices())
    d = {s : float("inf")  for s in G.vertices()}
    previous_vertices = { s: None for s in G.vertices()}
    d[e]=0
    while sousgraphe:
        current_vertex = min(sousgraphe, key=lambda vertex: d[vertex])
        if d[current_vertex] == float("inf"):
            break
            
        for neighbour in G.neighbors_out(current_vertex):
            print(current_vertex,neighbour)
            alternative_route = d[current_vertex] + G.get_edge_data(current_vertex,neighbour)['weight']
            if alternative_route < d[neighbour]:
                d[neighbour] = alternative_route
                previous_vertices[neighbour] = current_vertex
        sousgraphe.remove(current_vertex)
        
    return d        

```

2. Appliquez cette fonction au graphe ci-dessus, et en déduire la distance entre `A` et `J`.

```python
from graph import examples
G = examples.dijkstra()
G.show()
```

Output:
```
Figure(fig_margin={'top': 60, 'bottom': 60, 'left': 60, 'right': 60}, layout=Layout(height='400px', width='400…
```

```python
dijkstra(G, "A")
```

Output:
```
A B
A C
A E
B A
B F
F B
F I
E A
E J
C A
C G
C H
H C
H D
H J
G C
I F
I J
J E
J H
J I
D H
{'A': 0,
 'B': 87,
 'C': 217,
 'D': 503,
 'E': 173,
 'F': 167,
 'G': 403,
 'H': 320,
 'I': 417,
 'J': 487}
```

<!-- grade_id: cell-0d543769776d9f6a, points: 1, grade: true, solution: false -->
```python
assert dijkstra(G, "A")["A"] == 0
```

Output:
```
A B
A C
A E
B A
B F
F B
F I
E A
E J
C A
C G
C H
H C
H D
H J
G C
I F
I J
J E
J H
J I
D H
```

3. (Bonus) Adapter la fonction précédente pour qu'elle prenne deux sommets `e` et `f` et renvoie un plus court chemin entre `e` et `f`.

```python
from collections import deque
def dijkstra_bonus(G, e,dest):
    """
    Renvoie un dictionnaire associant à chaque sommet sa distance depuis `e`
    """
    sousgraphe=list(G.vertices())
    d = {s : float("inf")  for s in G.vertices()}
    previous_vertices = { s: None for s in G.vertices()}
    d[e]=0
    while sousgraphe:
        current_vertex = min(sousgraphe, key=lambda vertex: d[vertex])
        if d[current_vertex] == float("inf"):
            break
            
        for neighbour in G.neighbors_out(current_vertex):
            alternative_route = d[current_vertex] + G.get_edge_data(current_vertex,neighbour)['weight']
            if alternative_route < d[neighbour]:
                d[neighbour] = alternative_route
                previous_vertices[neighbour] = current_vertex
        sousgraphe.remove(current_vertex)
    current_vertex=dest
    path=deque()
    while previous_vertices[current_vertex] is not None:
            path.appendleft(current_vertex)
            current_vertex = previous_vertices[current_vertex]
    if path:
        path.appendleft(current_vertex)
    return path,d[path[-1]]


```

4. (Bonus) Instrumentez votre fonction pour en visualiser l'exécution.

<!-- grade_id: cell-06c0e69a490a9644, points: 1, grade: true, solution: true -->
```python
from collections import deque
def dijkstra_bonus_visu(G, e,dest):
    """
    Renvoie un dictionnaire associant à chaque sommet sa distance depuis `e`
    """
    sousgraphe=list(G.vertices())
    d = {s : float("inf")  for s in G.vertices()}
    previous_vertices = { s: None for s in G.vertices()}
    d[e]=0
    while sousgraphe:
        current_vertex = min(sousgraphe, key=lambda vertex: d[vertex])
        if d[current_vertex] == float("inf"):
            break
            
        for neighbour in G.neighbors_out(current_vertex):
            alternative_route = d[current_vertex] + G.get_edge_data(current_vertex,neighbour)['weight']
            if alternative_route < d[neighbour]:
                d[neighbour] = alternative_route
                previous_vertices[neighbour] = current_vertex
        sousgraphe.remove(current_vertex)
    current_vertex=dest
    path=deque()
    while previous_vertices[current_vertex] is not None:
            print(current_vertex)
            path.appendleft(current_vertex)
            current_vertex = previous_vertices[current_vertex]
    if path:
        path.appendleft(current_vertex)
    return path,d[path[-1]]


```

## Problème 1: Le chemin le plus rapide en métro de Montgallet à Billancourt ?
<center><img src="metro-paris.gif" width="50%"></center>

1. Le fichier [metro_complet.txt](metro_complet.txt) contient la description d'un graphe modélisant le métro de Paris. Consultez son contenu pour en comprendre le format.
2. Écrire une fonction qui lit le fichier et renvoie le graphe qu'il contient sous la forme d'un objet de type `Graph`.
3. Utilisez la fonction `dijkstra` pour calculer un plus court chemin de Montgallet à Billancourt!

<!-- grade_id: cell-7e01ead24612c65d, points: 3, grade: true, solution: true -->
```python
import pandas as pd
pd.read_csv("./metro_complet.txt",sep=" ",error_bad_lines=False,header=None)
```

Output:
```
b'Skipping line 4: expected 2 fields, saw 3\nSkipping line 5: expected 2 fields, saw 3\nSkipping line 7: expected 2 fields, saw 3\nSkipping line 10: expected 2 fields, saw 4\nSkipping line 11: expected 2 fields, saw 4\nSkipping line 12: expected 2 fields, saw 3\nSkipping line 13: expected 2 fields, saw 4\nSkipping line 14: expected 2 fields, saw 4\nSkipping line 16: expected 2 fields, saw 3\nSkipping line 17: expected 2 fields, saw 3\nSkipping line 18: expected 2 fields, saw 4\nSkipping line 22: expected 2 fields, saw 3\nSkipping line 27: expected 2 fields, saw 4\nSkipping line 31: expected 2 fields, saw 4\nSkipping line 32: expected 2 fields, saw 4\nSkipping line 35: expected 2 fields, saw 3\nSkipping line 36: expected 2 fields, saw 3\nSkipping line 39: expected 2 fields, saw 4\nSkipping line 40: expected 2 fields, saw 10\nSkipping line 44: expected 2 fields, saw 3\nSkipping line 50: expected 2 fields, saw 3\nSkipping line 51: expected 2 fields, saw 3\nSkipping line 52: expected 2 fie
```

```python
file1 = open('./metro_complet.txt', 'r') 

Lines = file1.read() 
df=pd.DataFrame()
l=Lines.split("\n")
cpt=0
while l:
    cpt+=1
    i=l.pop(0)
    if i == "coord sommets":
        break
    #print(i)
    df=df.append({"index":i.rsplit(" ")[0],"commune":i.rsplit(" ")[-1] },ignore_index=True)
cpt=0
df2=pd.DataFrame()
while l:
    cpt+=1
    i=l.pop(0)
    
    if i == "arcs values":
        #print(i)
        break
    #print(i)
    df2=df2.append({"index":i.split(" ")[0],"coord":i.split(" ")[1],"sommets":i.split(" ")[2] },ignore_index=True)
df3=pd.DataFrame()
while l:
    cpt+=1
    i=l.pop(0)
    if i =="":
        break
    
    #print(i)
    df3=df3.append({"sommet2":i.split(" ")[0],"sommet3":i.split(" ")[1],"value":i.split(" ")[2] },ignore_index=True)

    

#df3
df=df.drop([0,1])
df["index"]=df["index"].apply(lambda x : str(int(x)))
df
```

Output:
```
commune index
2          Abbesses     0
3             Dumas     1
4           Marceau     2
5            Alésia     3
6            France     4
..              ...   ...
373       Militaire   371
374  Maisons-Alfort   372
375       d'Auteuil   373
376          Pantin   374
377          Marcel   375

[376 rows x 2 columns]
```

```python
import numpy as np
df4=df3.copy()
df4["fusion"]=df4.apply(lambda x : sorted([x["sommet2"],x["sommet3"]]),axis=1)
df4["fusion"]=df4["fusion"].apply(lambda x : str(x[0])+" "+str(x[1]))
df4=df4.drop_duplicates(subset=["fusion"])
#df.loc[df["commune"]=="Fayette","index"].values[0]
#df.loc[df["commune"]=="Billancourt","index"].values[0]
edges=list(df4.apply(lambda x : (x["sommet2"],x["sommet3"],int(float(x["value"]))),axis=1))
vertices=list(df["index"])
```

```python
from graph_networkx import Graph, examples
G=Graph(vertices,edges=edges,directed=False)

```

```python
G.show()
```

Output:
```
Figure(fig_margin={'top': 60, 'bottom': 60, 'left': 60, 'right': 60}, layout=Layout(height='400px', width='400…
```

```python
def metro(depart,dest):
    
    dep=df.loc[df["commune"]==depart,"index"].values[0]
    des=df.loc[df["commune"]==dest,"index"].values[0]
    res=dijkstra_bonus(G, dep,des)
    l=[df.loc[df["index"]==i,"commune"].values[0] for i in res[0]]


    
    return l,res[1]

metro("Montgallet","Billancourt")
```

Output:
```
(['Montgallet',
  'Diderot',
  'Faidherbe-Chaligny',
  'Rollin',
  'Bastille',
  'Bastille',
  'Marais',
  'Ville',
  'Châtelet',
  'Rivoli',
  'Louvre',
  'Tuileries',
  'Concorde',
  'Clémenceau',
  'Roosevelt',
  'Roosevelt',
  'Marceau',
  'Iéna',
  'Trocadéro',
  'Pompe',
  'Muette',
  'Ranelagh',
  'Jasmin',
  'Auteuil',
  'Molitor',
  'Exelmans',
  'Saint-Cloud',
  'Sembat',
  'Billancourt'],
 1488)
```

```python
def metro_depart_plusieur(depart,dest):
    des=df.loc[df["commune"]==dest,"index"].values[0]
    dep=df.loc[df["commune"]==depart,"index"].values

    tmp= []
    tmp2=[]
    cpt=0
    for i in dep:
        print(i)
        
        res=dijkstra_bonus(G, i,des)
        tmp.append(res)
        tmp2.append(res[1])
        cpt+=1
    index = tmp2.index(min(tmp2))
    res=tmp[index]
    print(tmp2)
    
        
    
    l=[df.loc[df["index"]==i,"commune"].values[0] for i in res[0]]


    
    return l,res[1]

metro_depart_plusieur("Fayette","Billancourt")
```

Output:
```
59
60
[923, 803]
(['Fayette',
  'Caumartin',
  'Saint-Augustin',
  'Miromesnil',
  'Roule',
  'Roosevelt',
  'Marceau',
  'Iéna',
  'Trocadéro',
  'Pompe',
  'Muette',
  'Ranelagh',
  'Jasmin',
  'Auteuil',
  'Molitor',
  'Exelmans',
  'Saint-Cloud',
  'Sembat',
  'Billancourt'],
 803)
```

```python
def metro_depart_arrive_plusieur(depart,dest):
    des=df.loc[df["commune"]==dest,"index"].values
    dep=df.loc[df["commune"]==depart,"index"].values

    tmp= []
    tmp2=[]
    cpt=0
    for i in dep:
        for j in des:
        
            print(i,j)

            res=dijkstra_bonus(G, i,j)
            tmp.append(res)
            tmp2.append(res[1])
            cpt+=1
    index = tmp2.index(min(tmp2))
    res=tmp[index]
    print(tmp2)
    
        
    
    l=[df.loc[df["index"]==i,"commune"].values[0] for i in res[0]]


    
    return l,res[1]

metro_depart_arrive_plusieur("Fayette","Rochechouart")
```

Output:
```
59 13
59 14
60 13
60 14
[575, 455, 655, 535]
(['Fayette',
  'Peletier',
  'Cadet',
  'Poissonnière',
  "l'Est",
  "l'Est",
  'Nord',
  'Rochechouart'],
 455)
```

```python

```

```python

```