# Applications

## Composantes connexes
Soit $G$ un graphe non orienté. La composante connexe d'un sommet $s$ de $G$ est l'ensemble des sommets atteignables depuis $s$ en suivant un chemin dans $G$.

**Exercice**
1. Décrire un algorithme pour calculer toutes les composantes connexes d'un graphe non orienté
2. (Bonus) Implanter cet algorithme et visualiser son exécution

<!-- grade_id: cell-7edc9912d649c82e, points: 2, grade: true, solution: true -->
### 1°)
Le parcours en profondeur permet d'avoir une composante connexe d'un sommet s d'un graphe.
### 2°)


```python
from graph import Graph, examples
C3 = examples.dijkstra()
def explorer(graphe,sommet):
    marque[sommet] = 1
    print(sommet)
    for k in graphe.neighbors_out(sommet):
        if k not in marque:
            explorer(graphe,k)
    return marque

def Affcomposanteconnexe(graphe,sommet):
    global marque
    marque = {}
    for s in graphe.neighbors_out(sommet):
        if s not in marque:
            explorer(graphe,s)
Affcomposanteconnexe(C3, "D")

```

Output:
```
H
C
A
B
F
I
J
E
G
D
```

```python
C3.vertices()
```

Output:
```
('A', 'B', 'C', 'D', 'E', 'F', 'G', 'H', 'I', 'J')
```

```python
C3.show()
```

Output:
```
Figure(fig_margin={'top': 60, 'bottom': 60, 'left': 60, 'right': 60}, layout=Layout(height='400px', width='400…
```

## Généralisation

**Exercice**:

1. Généraliser la fonction `parcours` de la fiche [01-ParcoursDeGraphe.ipynb](01-ParcoursDeGraphe.ipynb) pour qu'elle prenne en argument la *fonction de voisinage* du graphe plutôt que le graphe lui-même.

<!-- grade_id: cell-203cbf9059dad7fa, points: 0, grade: false, solution: true -->
```python
def parcours_fonction(neighbors_out, u):
    marked = {u} # L'ensemble des sommets déjà rencontrés
    todo   = {u} # L'ensemble des sommets déjà rencontrés, mais pas encore traités
    
    while todo:
        # Invariants:
        # - Si `v` est dans `marked`, alors il y a un chemin de `u` à `v`
        # - Si `v` est dans `marked` et pas dans `todo`
        #   alors tous les voisins de `v` sont dans dans `marked`
        v = todo.pop()
        for w in neighbors_out(v):
            if w not in marked:
                marked.add(w)
                todo.add(w)
    return marked

```

2. Testez votre fonction sur l'exemple suivant.  
   **Indication**: étant donné un graphe `G` tel que implanté précédément, la fonction `voisins` requise par `distance` peut être construite comme suit:

```python
from graph import examples
G = examples.parcours_directed()
voisins = G.neighbors_out
```

<!-- grade_id: cell-c0385743a593b35c, points: 1, grade: true, solution: false -->
```python
assert parcours_fonction(voisins, 'D') == {'D', 'F', 'G', 'H'}
```

2. Reprenez de même tous les tests de la fiche 01 avec votre nouvelle fonction

```python
H = examples.disconnected()
H=H.neighbors_out
assert sorted(parcours_fonction(H, 1)) == [1, 2, 5]
assert sorted(parcours_fonction(H, 3)) == [3, 4]
H = examples.cours_1_G()
H=H.neighbors_out

assert sorted(parcours_fonction(H, 3)) == [0, 1, 2, 3, 4, 5]
H = examples.cours_1_reseau()
G = examples.parcours_directed()

G=G.neighbors_out
assert parcours_fonction(G, "B") == {'B', 'C', 'D', 'F', 'G', 'H'}
assert parcours_fonction(G, "A") == {'A', 'B', 'C', 'D', 'F', 'G', 'H'}
```

## Application: ensembles définis récursivement
De nombreux ensembles mathématiques sont définis en donnant quelques éléments initiaux, puis en applicant récursivement un processus permettant de produire de nouveaux éléments à partir d'éléments déjà présents. Par exemple, l'ensemble des entiers naturels est construit à partir de $0$ et de la fonction successeur $x \mapsto x+1$.

**Définition:**

Soit $E$ un ensemble, $R\subset E$ un sous ensemble fini, et $f$ une fonction associant à chaque élément de $E$ un sous-ensemble fini de $E$. L'ensemble $f$ défini récursivement par $R$ et $f$ est le plus petit sous-ensemble de $E$ contenant $R$ et *stable* par $f$: si $e\in F$ alors $f(e)\subset F$.

**Exercice**

Décrire les ensembles définis récursivement obtenus dans chacun des cas suivants:
- $E=\mathbb N$, $R=\{1\}$, $f(e) = \{ e + 2 \}$

e=\{1\} un sous ensemble de E<br>
f(1)=(1+2)= $3\subset E$<br>
f(3)=5<br>
.<br>
.<br>
.<br>
f(n)=n+2<br>
Represente nombre impair<br>
(Chemin des sommets de nombre impair.)



- $E=\mathbb N$, $R=\{1\}$, $f(e) = \{ 2e, e+3 \}$

```python
f(1)={2,4}
f(2)={4,5}----f(4)={8,7}
f(4)={8,7}----f(5)={10,8}----f(6)={12,9}----f(4)={8,7}
Cela represente les chiffre N a part les multiple de 3
```

Output:
```
File "<ipython-input-12-a3a557965a19>", line 4
    Cela represente les chiffre N a part les multiple de 3
                  ^
SyntaxError: invalid syntax
```

- $E$: listes, $R=\{()\}$, $f(u)$: rajouter $0$ ou $1$ à la fin du $u$

```python
f()={0ou 1}
f(0)={01ou 00}----f(1)={10 ou 01}
Cela represente toutes les combinaisons du code binaire.
```

Output:
```
File "<ipython-input-13-c5996a6877af>", line 1
    f()={0ou 1}
          ^
SyntaxError: invalid token
```

- $E$: listes, $R=\{(1,2,3,2,1)\}$, $f(u)$: supprimer la première ou dernière lettre de $u$

```python
def f(l):
    if(l==[]):
        return
    print(l)
    l1=l.copy()
    l2=l.copy()
    l1.pop(0)
    l2.pop(-1)
    f(l1)
    f(l2)

```

```python
f((1,2,3,2,1))={(2,3,2,1)ou(1,2,3,2) }
f((2,3,2,1))={(2,3,2)ou(3,2,1) }----f(1,2,3,2)={(2,3,2) ou (1,2,3)}
```

Output:
```
File "<ipython-input-15-c91ee2596e91>", line 1
    f((1,2,3,2,1))={(2,3,2,1)ou(1,2,3,2) }
                              ^
SyntaxError: invalid syntax
```

Cela represente les nombres (1,2,3,2,1) dans le meme ordre mais avec n'importe quel de ces nombres<br>
Ex:<br>
1.2.3.2.1<br>
1.2.3.2<br>
1.2.3<br>
1.2<br>
1<br>



```python
def f(l):
    res=[]
    for j in range(0,len(l)):
        if(j+1>=len(l)):
            break
        new=l.copy()
        new[j]+=new[j+1]
        new.pop(j+1)
        #print("new=",new,"res",res)
        if new not in res:

            res.append(new)
    return res
def frecur(l):
    #print(l,"\n\n\n")
    res=[]
    for i in l:
        #print("i=",i,"res",res,"fi",f(i))
        fi=f(i)
        der=fi[-1]
        j=0
        while j<len(fi):
            if fi[j] in res:
                #print("diiiiidifii",fi[j],res)
                fi.pop(j)
            j+=1
        

            
        res+=fi
    if len(res[0])==1:
        return res
    fi=frecur(res)
    j=0
    
    while j<len(fi):
        if fi[j] in res:
            fi.pop(j)
        j+=1
    res=res+fi
    d=[]
    i=0
    while i<len(res):
        if res[i] not in d:
            d.append(res[i])
        i+=1
        
    return d
def C_n(l):
    return l+frecur(l)
l=f([[1,1,1,1,1]])
C_n([[1,1,1,1,1,1]])
```

Output:
```
[[1, 1, 1, 1, 1, 1],
 [2, 1, 1, 1, 1],
 [1, 2, 1, 1, 1],
 [1, 1, 2, 1, 1],
 [1, 1, 1, 2, 1],
 [1, 1, 1, 1, 2],
 [3, 1, 1, 1],
 [2, 2, 1, 1],
 [2, 1, 2, 1],
 [2, 1, 1, 2],
 [1, 3, 1, 1],
 [1, 2, 2, 1],
 [1, 2, 1, 2],
 [1, 1, 3, 1],
 [1, 1, 2, 2],
 [1, 1, 1, 3],
 [4, 1, 1],
 [3, 2, 1],
 [3, 1, 2],
 [2, 3, 1],
 [2, 2, 2],
 [2, 1, 3],
 [1, 4, 1],
 [1, 3, 2],
 [1, 2, 3],
 [1, 1, 4],
 [5, 1],
 [4, 2],
 [3, 3],
 [2, 4],
 [1, 5],
 [6]]
```

**Exercice**

Soit $C_n$ l'ensemble récursivement défini par:
- $R=\{ (\underbrace{1,1,\cdots,1}_n)\}$
- $f$ qui a une liste associe toute les listes obtenues en regroupant et sommant deux entrées consécutives<br>
Par exemple, $f((2,3,1,1)) = \{(5,1,1), (2,4,1), (2,3,2)\}$

- Calculer à la main les **quatre** éléments de $C_3$, sans oublier $(4)$!
- Calculer avec la machine tous les éléments de $C_6$.
- Combien y en a t'il?

- Même chose pour $C_1,C_2,C_3,C_4,C_5$.
- Que conjecturez vous?
- Pouvez vous le prouver?

<!-- grade_id: cell-a16f2358ab641611, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Exercice** (bonus)
Une *permutation* est une liste d'entiers telle que, si $n$ est sa longueur, alors tous les entiers de $1$ à $n$ apparaissent exactement une fois. Décrire l'ensemble des permutations comme un ensemble énuméré. Utiliser cela pour lister toutes les permutations de longueur $n\leq 6$

<!-- grade_id: cell-629cf4c8130929e4, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

```python
1°)Calculer à la main les quatre éléments de  𝐶3 , sans oublier  (4) !
[[2, 1], [1, 2]]
[2, 1]=[3]
[1, 2]=[3]
=[[2, 1], [1, 2],[3],[3]]

2°)
[[1, 1, 1, 1, 1, 1],
 [2, 1, 1, 1, 1],
 [1, 2, 1, 1, 1],
 [1, 1, 2, 1, 1],
 [1, 1, 1, 2, 1],
 [1, 1, 1, 1, 2],
 [3, 1, 1, 1],
 [2, 2, 1, 1],
 [2, 1, 2, 1],
 [2, 1, 1, 2],
 [1, 3, 1, 1],
 [1, 2, 2, 1],
 [1, 2, 1, 2],
 [1, 1, 3, 1],
 [1, 1, 2, 2],
 [1, 1, 1, 3],
 [4, 1, 1],
 [3, 2, 1],
 [3, 1, 2],
 [2, 3, 1],
 [2, 2, 2],
 [2, 1, 3],
 [1, 4, 1],
 [1, 3, 2],
 [1, 2, 3],
 [1, 1, 4],
 [5, 1],
 [4, 2],
 [3, 3],
 [2, 4],
 [1, 5],
 [6]]
```

Output:
```
File "<ipython-input-17-0cc1b1640a16>", line 1
    1°)Calculer à la main les quatre éléments de  𝐶3 , sans oublier  (4) !
     ^
SyntaxError: invalid character in identifier
```

```python
3°) Il y en a 32 en comptant les doublons
```

Output:
```
File "<ipython-input-18-20d12d8e9684>", line 1
    3°) Il y en a 32 en comptant les doublons
     ^
SyntaxError: invalid character in identifier
```

```python
#4°)
for i in range(2,6):
    print("i=",i,"len=",len(C_n([[1]*i])))
```

Output:
```
i= 2 len= 2
i= 3 len= 4
i= 4 len= 8
i= 5 len= 16
```

```python
#5°)p(n)=Que cela tend toujours vers une liste [n]
```

#6°)
Demontrons par recurrence:<br>
p(n)=[n]<br>
[]
n=0<br>
Somme liste vide =0<br>
HR= pn vrai pour un certain n<br>
Montrons pour pn+1<br>

Heridité:<br>
f(1,...,1,1)<br>
par HR f(1,...1,1) des n premier 1 =n<br>
donc f(1,...1,1)=f(n,1)par HR<br>
Par definition de f nous avons f(n,1)=[n+1] <br>
D'ou P(n+1) vrai =[n+1]<br>

Conclusion:<br>
Par récurrence, la propriété est vraie pour tout n.<br>




```python

```

```python

```

```python

```