# Problème 2: [Rush Hour](http://www.thinkfun.com/products/rush-hour/)

La voiture rouge est coincée dans un embouteillage; comment déplacer les véhicules pour la faire sortir?
Ci-dessous le défi 1.
<center><img src="rush_hour.gif" width="30%"></center>

Quel rapport avec les parcours de graphes? Saurez-vous faire résoudre par l'ordinateur le défi 40 en un nombre minimum du coups?

Dans cette feuille, vous découvrirez le modèle et deux mini applications interactives pour le jeu RushHour qui vous sont fournis. Il ne restera «plus qu'à» implanter l'«Intelligence Artificielle» pour résoudre le jeu.

## Un modèle pour RushHour

### Quelques exemple d'utilisation programmatique

```python
from rush_hour import Plateau
```

```python
plateau = Plateau(['A2R00','X2R21','C2R44','R3R52','O3D05','P3D10','Q3D13','B2D40']); plateau
```

Output:
```
+------+
|AA   O|
|P  Q O|
|PXXQ O 
|P  Q  |
|B   CC|
|B RRR |
+------+
```

```python
plateau.recule("R")
```

Output:
```
+------+
|AA   O|
|P  Q O|
|PXXQ O 
|P  Q  |
|B   CC|
|BRRR  |
+------+
```

```python
plateau.est_gagnant()
```

Output:
```
False
```

Chargement du 40ème défi de RushHour (voir les fichiers dans <a href="RushHourDefis/">RushHourDefis/</a>):

```python
Plateau(1)
```

Output:
```
+------+
|AA   O|
|P  Q O|
|PXXQ O 
|P  Q  |
|B   CC|
|B RRR |
+------+
```

```python
    def solution(niveau, return_graph=False, verbeux=False):
        '''
        Renvoie une solution optimale pour ce niveau

            >>> RushHour.solution(1) # doctest: +SKIP
            ['AR1', 'CL3', 'OD3', 'PU1', 'BU1', 'RL2', 'QD2', 'XR3']
        '''
        ### BEGIN SOLUTION
        if verbeux:
            progressbar = tqdm.tqdm() # tqdm.tqdm_notebook()
        plateau = Plateau(niveau)
        afaire = deque([(plateau, 0)])
        predecesseurs = {plateau: None}
        if plateau.est_gagnant():
            return []
        olddist = 0
        while afaire:
            sommet, dist = afaire.popleft()
            for label, voisin in sommet.voisins().items():
                if voisin in predecesseurs:
                    continue
                if verbeux:
                    progressbar.update()
                    progressbar.set_postfix({'afaire': len(afaire)})
                predecesseurs[voisin] = label, sommet
                afaire.append((voisin, dist+1))
                if voisin.est_gagnant():
                    if return_graph:
                        return predecesseurs
                    sommet = voisin
                    chemin = []
                    while predecesseurs[sommet] != None:
                        label, sommet = predecesseurs[sommet]
                        chemin.append(label)
                    chemin.reverse()
                    return chemin
            if dist != olddist:
                olddist = dist
        if return_graph:
            return predecesseurs
        return None
        ### END SOLUTION

```

```python
def parcours_en_largeur(plateau, u):
    """
    INPUT:
    - 'G' - un graphe
    - 'u' - un sommet du graphe
    
    OUTPUT: un dictionnaire associant à chaque sommet `v` accessible depuis `u` sa distance depuis `u`
    """
    distances = {u: 0} # L'ensemble des sommets déjà rencontrés
    todo      = [u]    # Une liste de sommets à traiter

    while todo:
                
        v = todo.pop()
        plateau=plateau.deplace(v)
        possible=plateau.voisins()
        for w in possible:
            if w not in distances:
                distances[w] = distances[v] + 1
                todo.append(w)

        # Invariants:
        # - Si `v` est dans `distance`, alors il y a un chemin de `u` à `v`,
        #   et distance[v] contient la distance de `u` à `v`;
        # - Si `v` est dans `distance` et pas dans `todo`
        #   alors tous les voisins de `v` sont dans dans `distance`
        # YOUR CODE HERE
    return distances
```

### Un exemple d'interface utilisateur minimale

```python
from ipywidgets import interact_manual
from rush_hour import Plateau
plateau = Plateau(['A2R00','X2R21','C2R44','R3R52','O3D05','P3D10','Q3D13','B2D40']); plateau
@interact_manual
def step(voiture=plateau.voitures.keys(), distance=[0,-1,1]):
    global plateau
    plateau2 = plateau.avance(voiture, distance)
    if plateau2:
        plateau = plateau2
    return plateau
```

Output:
```
interactive(children=(Dropdown(description='voiture', options=('A', 'X', 'C', 'R', 'O', 'P', 'Q', 'B'), value=…
```

## Une application graphique basée sur les widgets de Jupyter

```python
from rush_hour_application import RushHourApplication
A = RushHourApplication()
A
```

Output:
```
RushHourApplication(children=(HBox(children=(Dropdown(options=(('Niveau 1', 1), ('Niveau 3', 3), ('Niveau 7', …
```

## À vous de jouer!

Implantez la méthode `solution` dans la class [`RushHour`](rush_hour.py) pour déterminer une séquence minimale de coup permettant de résoudre un défi, puis utilisez la ci-dessous pour résoudre le défi 40 de Rush Hour.

**Indication**: considérez le graphe dont les sommets sont les états possibles du plateau et les arêtes les déplacements de voiture!

```python
from rush_hour import RushHour
solution = RushHour.solution(40)
```

<!-- grade_id: cell-38c9c48e258b9fd9, points: 2, grade: true, solution: false -->
```python
assert RushHour.est_solution(40, solution)
```

```python
RushHour.est_solution(40, solution)
```

Output:
```
True
```

```python

```